Flamingo participates in multiple models of cell competition

Author response:

The following is the authors’ response to the original reviews.

Reviewer 1:

Summary:

This paper is focused on the role of Cadherin Flamingo (Fmi) - also called Starry night (stan) - in cell competition in developing Drosophila tissues. A primary genetic tool is monitoring tissue overgrowths caused by making clones in the eye disc that express activated Ras (RasV12) and that are depleted for the polarity gene scribble (scrib). The main system that they use is ey-flp, which makes continuous clones in the developing eye-antennal disc beginning at the earliest stages of disc development. It should be noted that RasV12, scrib-i (or lgl-i) clones only lead to tumors/overgrowths when generated by continuous clones, which presumably creates a privileged environment that insulates them from competition. Discrete (hs-flp) RasV12, lgl-i clones are in fact outcompeted (PMID: 20679206), which is something to bear in mind.

We think it is unlikely that the outcome of RasV12, scrib (or lgl) competition depends on discrete vs. continuous clones or on creation of a privileged environment. As shown in the same reference mentioned by the reviewer, the outcome of RasV12, scrib (or lgl) tumors greatly depends on the clone being able to grow to a certain size. The authors show instances of discrete clones where larger RasV12, lgl clones outcompete the surrounding tissue and eliminate WT cells by apoptosis, whereas smaller clones behave more like losers. It is not clear what aspect of the environment determines the ability of some clones to grow larger than others, but in neither case are the clones prevented from competition. Other studies show that in mammalian cells, RasV12, scrib clones are capable of outcompeting the surrounding tissue, such as in Kohashi et al (2021), where cells carrying both mutations actively eliminate their neighbors.

The authors show that clonal loss of Fmi by an allele or by RNAi in the RasV12, scrib-i tumors suppresses their growth in both the eye disc (continuous clones) and wing disc (discrete clones). The authors attributed this result to less killing of WT neighbors when Myc over-expressing clones lacking Fmi, but another interpretation (that Fmi regulates clonal growth) is equally as plausible with the current results.

See point (1) for a discussion on this.

Next, the authors show that scrib-RNAi clones that are normally out-competed by WT cells prior to adult stages are present in higher numbers when WT cells are depleted for Fmi. They then examine death in RasV12, scrib-i ey-FLP clones, or in discrete hsFLP UAS-Myc clones. They state that they see death in WT cells neighboring RasV12, scrib-i clones in the eye disc (Figures 4A-C). Next, they write that RasV12, scrib-I cells become losers (i.e., have apoptosis markers) when Fmi is removed. Neither of these results are quantified and thus are not compelling. They state that a similar result is observed for Myc over-expression clones that lack Fmi, but the image was not compelling, the results are not quantified and the controls are missing (Myc over-expressing clones alone and Fmi clones alone).

We assayed apoptosis in UAS-Myc clones in eye discs but neglected to include the results in Figure 4. We include them in the updated manuscript. Regarding Fmi clones alone, we direct the reviewer’s attention to Fig. 2 Supplement 1 where we showed that fminull clones cause no competition. Dcp-1 staining showed low levels of apoptosis unrelated to the fminull clones or twin-spots.

Regarding the quantification of apoptosis, we did not provide a quantification, in part because we observe a very clear visual difference between groups (Fig. 4A-K), and in part because it is challenging to come up with a rigorous quantification method. For example, how far from a winner clone can an apoptotic cell be and still be considered responsive to the clone? For UASMyc winner clones, we observe a modest amount of cell death both inside and outside the clones, consistent with prior observations. For fminull UAS-Myc clones, we observe vastly more cell death within the fminull UAS-Myc clones and modest death in nearby wildtype cells, and consequently a much higher ratio of cell death inside vs outside the clone. Because of the somewhat arbitrary nature of quantification, and the dramatic difference, we initially chose not to provide a quantification. However, given the request, we chose an arbitrary distance from the clone boundary in which to consider dying cells and counted the numbers for each condition. We view this as a very soft quantification, but we nevertheless report it in a way that captures the phenomenon in the revised manuscript.

They then want to test whether Myc over-expressing clones have more proliferation. They show an image of a wing disc that has many small Myc overexpressing clones with and without Fmi. The pHH3 results support their conclusion that Myc overexpressing clones have more pHH3, but I have reservations about the many clones in these panels (Figures 5L-N).

As the reviewer’s reservations are not specified, we have no specific response.

They show that the cell competition roles of Fmi are not shared by another PCP component and are not due to the Cadherin domain of Fmi. The authors appear to interpret their results as Fmi is required for winner status. Overall, some of these results are potentially interesting and at least partially supported by the data, but others are not supported by the data.

Strengths:

Fmi has been studied for its role in planar cell polarity, and its potential role in competition is interesting.

Weaknesses:

(1) In the Myc over-expression experiments, the increased size of the Myc clones could be because they divide faster (but don't outcompete WT neighbors). If the authors want to conclude that the bigger size of the Myc clones is due to out-competition of WT neighbors, they should measure cell death across many discs of with these clones. They should also assess if reducing apoptosis (like using one copy of the H99 deficiency that removes hid, rpr, and grim) suppresses winner clone size. If cell death is not addressed experimentally and quantified rigorously, then their results could be explained by faster division of Myc over-expressing clones (and not death of neighbors). This could also apply to the RasV12, scrib-i results.

Indeed, Myc clones have been shown to divide faster than WT neighbors, but that is not the only reason clones are bigger. As shown in (de la Cova et al, 2004), Myc-overexpressing cells induce apoptosis in WT neighbors, and blocking this apoptosis results in larger wings due to increased presence of WT cells. Also, (Moreno and Basler, 2004) showed that Myc-overexpressing clones cause a reduction in WT clone size, as WT twin spots adjacent to 4xMyc clones are significantly smaller than WT twin spots adjacent to WT clones. In the same work, they show complete elimination of WT clones generated in a tub-Myc background. Since then, multiple papers have shown these same results. It is well established then that increased cell proliferation transforms Myc clones into supercompetitors and that in the absence of cell competition, Myc-overexpressing discs produce instead wings larger than usual.

In (de la Cova et al, 2004) the authors already showed that blocking apoptosis with H99 hinders competition and causes wings with Myc clones to be larger than those where apoptosis wasn’t blocked. As these results are well established from prior literature, there is no need to repeat them here.

(2) This same comment about Fmi affecting clone growth should be considered in the scrib RNAi clones in Figure 3.

In later stages, scrib RNAi clones in the eye are eliminated by WT cells. While scrib RNAi clones are not substantially smaller in third instar when competing against fmi cells (Fig 3M), by adulthood we see that WT clones lacking Fmi have failed to remove scrib clones, unlike WT clones that have completely eliminated the scrib RNAi clones by this time. We therefore disagree that the only effect of Fmi could be related to rate of cell division.

(3) I don't understand why the quantifications of clone areas in Figures 2D, 2H, 6D are log values. The simple ratio of GFP/RFP should be shown. Additionally, in some of the samples (e.g., fmiE59 >> Myc, only 5 discs and fmiE59 vs >Myc only 4 discs are quantified but other samples have more than 10 discs). I suggest that the authors increase the number of discs that they count in each genotype to at least 20 and then standardize this number.

Log(ratio) values are easier to interpret than a linear scale. If represented linearly, 1 means equal ratios of A and B, while 2A/B is 2 and A/2B is 0.5. And the higher the ratio difference between A and B, the starker this effect becomes, making a linear scale deceiving to the eye, especially when decreased ratios are shown. Using log(ratios), a value of 0 means equal ratios, and increased and decreased ratios deviate equally from 0.

Statistically, either analyzing a standardized number of discs for all conditions or a variable number not determined beforehand has no effect on the p-value, as long as the variable n number is not manipulated by p-hacking techniques, such as increasing the n of samples until a significant p-value has been obtained. While some of our groups have lower numbers, all statistical analyses were performed after all samples were collected. For all results obtained by cell counts, all samples had a minimum of 10 discs due to the inherent though modest variability of our automated cell counts, and we analyzed all the discs that we obtained from a given experiment, never “cherry-picking” examples. For the sake of transparency, all our graphs show individual values in addition to the distributions so that the reader knows the n values at a glance.

(5) Figure 4 - shows examples of cell death. Cas3 is written on the figure but Dcp-1 is written in the results. Which antibody was used? The authors need to quantify these results. They also need to show that the death of cells is part of the phenotype, like an H99 deficiency, etc (see above).

Thank you for flagging this error. We used cleaved Dcp-1 staining to detect cell death, not Cas3 (Drice in Drosophila). We updated all panels replacing Cas3 by Dcp-1.

As described above, cell death is a well established consequence of myc overexpression induced cell death and we feel there is no need to repeat that result. To what extent loss of Fmi induces excess cell death or reduces proliferation in “would-be” winners, and to what extent it reduces “would-be” winners’ ability to eliminate competitors are interesting mechanistic questions that are beyond the scope of the current manuscript.

(6) It is well established that clones overexpressing Myc have increased cell death. The authors should consider this when interpreting their results.

We are aware that Myc-overexpressing clones have increased cell death, but it has also been demonstrated that despite that fact, they behave as winners and eliminate WT neighboring cells. And as mentioned in comment (1), WT clones generated in a 3x and 4x Myc background are eliminated and removed from the tissue, and blocking cell death increases the size of WT “losers” clones adjacent to Myc overexpressing clones.

(7) A better characterization of discrete Fmi clones would also be helpful. I suggest inducing hs-flp clones in the eye or wing disc and then determining clone size vs twin spot size and also examining cell death etc. If such experiments have already been done and published, the authors should include a description of such work in the preprint.

We have already analyzed the size of discrete Fmi clones and showed that they did not cause any competition, with fmi-null clones having the same size as WT clones in both eye and wing discs. We direct the reviewer’s attention to Figure 2 Supplement 1.

(8) We need more information about the expression pattern of Fmi. Is it expressed in all cells in imaginal discs? Are there any patterns of expression during larval and pupal development?

Fmi is equally expressed by all cells in all imaginal discs in Drosophila larva and pupa. We include this information and the relevant reference (Brown et al, 2014) in the updated manuscript.

(9) Overall, the paper is written for specialists who work in cell competition and is fairly difficult to follow, and I suggest re-writing the results to make it accessible to a broader audience.

We have endeavored to both provide an accessible narrative and also describe in sufficient detail the data from multiple models of competition and complex genetic systems. We hope that most readers will be able, at a minimum, to follow our interpretations and the key takeaways, while those wishing to examine the nuts and bolts of the argument will find what they need presented as simply as possible.

Reviewer 2:

Summary:

In this manuscript, Bosch et al. reveal Flamingo (Fmi), a planar cell polarity (PCP) protein, is essential for maintaining 'winner' cells in cell competition, using Drosophila imaginal epithelia as a model. They argue that tumor growth induced by scrib-RNAi and RasV12 competition is slowed by Fmi depletion. This effect is unique to Fmi, not seen with other PCP proteins. Additional cell competition models are applied to further confirm Fmi's role in 'winner' cells. The authors also show that Fmi's role in cell competition is separate from its function in PCP formation.

We would like to thank the reviewer for their thoughtful and positive review.

Strengths:

(1) The identification of Fmi as a potential regulator of cell competition under various conditions is interesting.

(2) The authors demonstrate that the involvement of Fmi in cell competition is distinct from its role in planar cell polarity (PCP) development.

Weaknesses:

(1) The authors provide a superficial description of the related phenotypes, lacking a comprehensive mechanistic understanding. Induction of apoptosis and JNK activation are general outcomes, but it is important to determine how they are specifically induced in Fmi-depleted clones. The authors should take advantage of the power of fly genetics and conduct a series of genetic epistasis analyses.

We appreciate that this manuscript does not address the mechanism by which Fmi participates in cell competition. Our intent here is to demonstrate that Fmi is a key contributor to competition. We indeed aim to delve into mechanism, are currently directing our efforts to exploring how Fmi regulates competition, but the size of the project and required experiments are outside of the scope of this manuscript. We feel that our current findings are sufficiently valuable to merit sharing while we continue to investigate the mechanism linking Fmi to competition.

(2) The depletion of Fmi may not have had a significant impact on cell competition; instead, it is more likely to have solely facilitated the induction of apoptosis.

We respectfully disagree for several reasons. First, loss of Fmi is specific to winners; loss of Fmi has no effect on its own or in losers when confronting winners in competition. And in the Ras V12 tumor model, loss of Fmi did not perturb whole eye tumors – it only impaired tumor growth when tumors were confronted with competitors. We agree that induction of apoptosis is affected, but so too is proliferation, and only when in winners in competition.

(3) To make a solid conclusion for Figure 1, the authors should investigate whether complete removal of Fmi by a mutant allele affects tumor growth induced by expressing RasV12 and scrib RNAi throughout the eye.

We agree with the reviewer that this is a worthwhile experiment, given that RNAi has its limitations. However, as fmi is homozygous lethal at the embryo stage, one cannot create whole disc tumors mutant for fmi. As an approximation to this condition, we have introduced the GMR-Hid, cell-lethal combination to eliminate non-tumor tissue in the eye disc. Following elimination of non-tumor cells, there remains essentially a whole disc harboring fminull tumor. Indeed, this shows that whole fminull tumors overgrow similar to control tumors, confirming that the lack of Fmi only affects clonal tumors. We provide those results in the updated manuscript (Figure 1 Suppl 2 C-D).

(4) The authors should test whether the expression level of Fmi (both mRNA and protein) changes during tumorigenesis and cell competition.

This is an intriguing point that we considered worthwhile to examine. We performed immunostaining for Fmi in clones to determine whether its levels change during competition. Fmi is expressed ubiquitously at apical plasma membranes throughout the disc, and this was unchanged by competition, including inside >>Myc clones and at the clone boundary, where competition is actively happening. We provide these results as a new supplementary figure (Figure 5 Suppl 1) in the updated manuscript.

Reviewer 3:

Summary:

In this manuscript, Bosch and colleagues describe an unexpected function of Flamingo, a core component of the planar cell polarity pathway, in cell competition in the Drosophila wing and eye disc. While Flamingo depletion has no impact on tumour growth (upon induction of Ras and depletion of Scribble throughout the eye disc), and no impact when depleted in WT cells, it specifically tunes down winner clone expansion in various genetic contexts, including the overexpression of Myc, the combination of Scribble depletion with activation of Ras in clones or the early clonal depletion of Scribble in eye disc. Flamingo depletion reduces the proliferation rate and increases the rate of apoptosis in the winner clones, hence reducing their competitiveness up to forcing their full elimination (hence becoming now "loser"). This function of Flamingo in cell competition is specific to Flamingo as it cannot be recapitulated with other components of the PCP pathway, and does not rely on the interaction of Flamingo in trans, nor on the presence of its cadherin domain. Thus, this function is likely to rely on a non-canonical function of Flamingo which may rely on downstream GPCR signaling.

This unexpected function of Flamingo is by itself very interesting. In the framework of cell competition, these results are also important as they describe, to my knowledge, one of the only genetic conditions that specifically affect the winner cells without any impact when depleted in the loser cells. Moreover, Flamingo does not just suppress the competitive advantage of winner clones, but even turns them into putative losers. This specificity, while not clearly understood at this stage, opens a lot of exciting mechanistic questions, but also a very interesting long-term avenue for therapeutic purposes as targeting Flamingo should then affect very specifically the putative winner/oncogenic clones without any impact in WT cells.

The data and the demonstration are very clean and compelling, with all the appropriate controls, proper quantification, and backed-up by observations in various tissues and genetic backgrounds. I don't see any weakness in the demonstration and all the points raised and claimed by the authors are all very well substantiated by the data. As such, I don't have any suggestions to reinforce the demonstration.

While not necessary for the demonstration, documenting the subcellular localisation and levels of Flamingo in these different competition scenarios may have been relevant and provided some hints on the putative mechanism (specifically by comparing its localisation in winner and loser cells).

Also, on a more interpretative note, the absence of the impact of Flamingo depletion on JNK activation does not exclude some interesting genetic interactions. JNK output can be very contextual (for instance depending on Hippo pathway status), and it would be interesting in the future to check if Flamingo depletion could somehow alter the effect of JNK in the winner cells and promote downstream activation of apoptosis (which might normally be suppressed). It would be interesting to check if Flamingo depletion could have an impact in other contexts involving JNK activation or upon mild activation of JNK in clones.

We would like to thank the reviewer for their thorough and positive review.

Strengths:

- A clean and compelling demonstration of the function of Flamingo in winner cells during cell competition.

- One of the rare genetic conditions that affects very specifically winner cells without any impact on losers, and then can completely switch the outcome of competition (which opens an interesting therapeutic perspective in the long term)

Weaknesses:

- The mechanistic understanding obviously remains quite limited at this stage especially since the signaling does not go through the PCP pathway.

Reviewer 2 made the same comment in their weakness (1), and we refer to that response. In future work, we are excited to better understand the pathways linking Fmi and competition.

Changes from the earlier version to the revised version

Abstract

Abstract

The growth and survival of cells with different fitness, such as those with a proliferative advantage or a deleterious mutation, is controlled through cell competition. During development, cell competition enables healthy cells to eliminate less fit cells that could jeopardize tissue integrity, and facilitates the elimination of pre-malignant cells by healthy cells as a surveillance mechanism to prevent oncogenesis. Malignant cells also benefit from cell competition to promote their expansion. Despite its ubiquitous presence, the mechanisms governing cell competition, particularly those common to developmental competition and tumorigenesis, are poorly understood. Here, we show that in Drosophila, the planar cell polarity (PCP) protein Flamingo (Fmi) is required by winners to maintain their status during cell competition in malignant tumors to overtake healthy tissue, in early pre-malignant cells when they overproliferate among wildtype cells, in healthy cells when they later eliminate pre-malignant cells, and by supercompetitors as they compete to occupy excessive territory within wildtype tissues. “Would-be” winners that lack Fmi are unable to over-proliferate, and instead become losers. We demonstrate that the role of Fmi in cell competition is independent of PCP, and that it uses a distinct mechanism that may more closely resemble one used in other less well-defined functions of Fmi.

Introduction

Dividing cells in proliferative tissues must maintain tissue homeostasis and structural integrity as well as preserve their genetic integrity. One of the mechanisms operating in proliferative tissues to achieve these requirements is cell competition. Cell competition is a process in which cells of a higher fitness (“winners”) eliminate less fit neighbors (“losers”), by inducing cell death (Norman et al., 2012; Meyer et al., 2014; Di Giacomo et al., 2017; Watanabe et al., 2018) and/or extrusion from the epithelial layer (Norman et al., 2012; Kon et al., 2017; Kohashi et al., 2021). Cell competition was first recognized in 1975 (Morata & Ripoll, 1975), when Morata and Ripoll studied the growth of Drosophila carrying a mutation in the gene for the ribosomal protein RpS17 (termed minute). Animals heterozygous for this minute mutation are viable but develop slowly due to decreased proliferation rates. They further observed that clones of cells carrying this minute mutation in a wildtype background were eliminated from the Drosophila developing wing. Since then, extensive advances have been made in our understanding of the process (reviewed in (Morata, 2021)). We now know that cell competition can be triggered by mutations that cause a disadvantage, such as the aforementioned minute cells, and can also be triggered by mutations that confer a proliferative advantage, such as ectopic expression of proto-oncogenes such as Myc or Ras (Moreno & Basler, 2004; de la Cova et al., 2004; Hogan et al., 2009). In this case, the process is known as super-competition, and wildtype cells behave as losers and are eliminated from the tissue by the mutant super-competitors. The “survival of the fittest” effect observed in cell competition has been shown to be critical during embryonic and larval development as well as for tissue homeostasis (Amoyel & Bach, 2014; Maruyama & Fujita, 2017; Morata, 2021). Much additional research has shown that cells and their neighbors are constantly evaluating their fitness, and that the active elimination of less fit cells is critical to maintain tissue integrity (Ellis et al., 2019), maintain chromosomal stability (Baillon et al., 2018; Ji et al., 2021), ameliorate effects of cellular aging (Merino et al., 2015) and suppress tumorigenesis (Kon et al., 2017; de Vreede et al., 2022).

While cell competition has been observed in many cell types, much attention has focused on studies in epithelia (Vincent et al., 2013). Polarity is a hallmark of epithelial tissues, and loss of polarity plays a prominent role in triggering cell competition and in the process of eliminating loser cells. Cells harboring mutations in apicobasal polarity genes have been shown to induce and/or regulate cell competition when surrounded by wildtype cells. In Drosophila, mutations in apicobasal polarity genes scribble (scrib), discs large (dlg), and lethal giant larva (lgl) all trigger cell competition when induced clonally (i.e. surrounded by wildtype cells), and the clones are eliminated from the epithelial tissue (Hariharan & Bilder, 2006; Morata & Calleja, 2020). However, clones that are deficient for any of these polarity genes can survive if they are competing against less fit cells, or if they carry additional mutations that confer a growth advantage. For example, when lgl clones, which would be eliminated when surrounded by wildtype cells, are made to express elevated levels of Myc, the competition is reversed, and they instead become winners (Froldi et al., 2010). This context-dependence and fitness surveillance highlights the complexity of cell competition and its plasticity.

In epithelial tissues, cell competition is critical to suppress tumorigenesis, as polarity-deficient clones are prone to malignancy. When scrib mutant epithelial cells acquire elevated activity of a proto-oncogene such as the constitutively active RasV12, these clones acquire malignant properties of ectopic growth and invasiveness (Pagliarini & Xu, 2003; Brumby & Richardson, 2003). Surveillance and rapid elimination of cells that develop polarity defects therefore serves to protect epithelia from oncogenesis. This competitive mechanism is known as Epithelial Defense Against Cancer, or EDAC (Maruyama & Fujita, 2017; Kon & Fujita, 2021). Through EDAC, normal epithelial cells eliminate neighboring transformed cells by inducing apoptosis or by inducing their apical or basal extrusion from the epithelial layer (Watanabe et al., 2018; Kohashi et al., 2021). On the other hand, tumors that exhibit more aggressive properties and escape the defense mechanisms not only proliferate but actively engage in cell competition to promote their own growth and malignancy (Vishwakarma & Piddini, 2020; Madan et al., 2022). For example, growth of Drosophila APC-/- intestinal adenomas requires active elimination of wildtype cells, and inhibiting cell death in the wildtype tissue hampers adenoma growth (Suijkerbuijk et al., 2016).

In addition to apicobasal polarity, epithelial tissues are planar polarized. Planar cell polarity (PCP) signaling polarizes cells within the plane of the epithelium to orient cellular structures, cell divisions, and cell migration during development and homeostasis (Adler, 2002; Simons & Mlodzik, 2008; Vladar et al., 2009; Butler & Wallingford, 2017). Intercellular Fmi homodimers scaffold the assembly of a set of core PCP proteins: on one side of a cell, Frizzled (Fz), Dishevelled (Dsh), and Diego (Dgo) comprise a complex that interacts with another complex containing Van Gogh (Vang) and Prickle (Pk) on the opposite side of the adjacent cell. Fmi homodimers transmit differential signals in opposite directions to communicate the presence of either protein complex to the other, linking the proximal and distal complexes and mediating asymmetric signaling between them (Lawrence et al., 2004; Strutt & Strutt, 2007; Strutt & Strutt, 2008; Chen et al., 2008; Struhl et al., 2012). In addition to establishing planar polarity, numerous reports have suggested links between PCP signaling and cancer progression, promoting cell motility, invasiveness and metastasis (Weeraratna et al., 2002; Katoh, 2005; Katoh & Katoh, 2007; Coyle et al., 2008; Gujral et al., 2014; VanderVorst et al., 2018; VanderVorst et al., 2023). Its roles, however, remain poorly characterized.

To probe a potential connection between PCP and cancer, we employed a Drosophila eye tumor model. We discovered a requirement for Fmi, but not other core PCP proteins, in tumor associated cell competition. We found that aggressive tumors require Fmi to outcompete the neighboring wildtype tissue. We then found that this Fmi requirement is not unique to tumor competition; in several developmental cell competition scenarios, removing Fmi from winner cells prevents them from eliminating their neighbors and transforms them into losers. “Would be” winners that lack Fmi show both a decrease in proliferation and an increase in apoptosis. By several criteria, we show that this function for Fmi is independent of its role in PCP signaling.

Results

Flamingo is required for tumorigenesis in Drosophila RasV12 tumors

Planar polarity has been linked to tumorigenesis in several organisms, tumor models, and patient tumor samples (Kaucká et al., 2013; Asad et al., 2014; Puvirajesinghe et al., 2016; VanderVorst et al., 2018; Chen et al., 2021c; Li et al., 2021; Chen et al., 2021b). We explored how the core PCP complex might be involved in tumorigenesis. To do so, we used the versatile and widely used RasV12 tumor model in Drosophila eye imaginal discs. In this model, eye disc cells are transformed into highly tumorigenic cells by co-expressing the constitutively active RasV12 oncoprotein and RNAi against the tumor suppressor Scribble (Scrib) (Pagliarini & Xu, 2003). Cells expressing these genes quickly expand into massive tumors and invade the neighboring brain tissue (Fig 1A-C).

Added in the revised version: Fmi is required in clonal tumors to outcompete wildtype tissue. A) Schematic of the whole eye disc RasV12, scrib RNAi, RFP tumors. B-H) Pupal RFP tumors imaged through the cuticle. All pupae are ey-Flp; act5C>CD2>Gal4, UAS-RFP with the indicated RNAi. C) Control UAS-RasV12, UAS-scrib RNAi, UAS-w RNAi tumors. D-H) UAS-RasV12, UAS-scrib RNAi tumors co-expressing UAS-RNAi against the PCP genes indicated above each pupa. B’-H’) Representative third instar larval brain and eye discs for each of the experimental groups from above. Arrowheads in B’ point to the eye-antenna imaginal disc (orange) and the brain lobes (green). I) Schematic of eye disc RasV12, scrib RNAi, RFP clonal tumors. J-P) Third instar larval eye discs RFP tumor clones generated via ey-Flp; FRT42D Gal80 / FRT42D; act5C>CD2>Gal4, UAS-RFP. J) Eye discs with non-tumor, control RFP clones. K) Control RasV12, UAS-scrib RNAi, RFP clonal tumors. L-P) RasV12, scrib RNAi, RFP clonal tumors carrying the PCP allele indicated above each panel. Scalebars: 100 μm

Removed in the revised version: Fmi is required in clonal tumors to outcompete wildtype tissue.

Removed in the revised version: A) Schematic of the whole eye disc RasV12, scrib RNAi, RFP tumors.

Removed in the revised version: B-H) Pupal RFP tumors imaged through the cuticle. All pupae are ey-Flp ; act5C>CD2>Gal4, UAS-RFP with the indicated RNAi.

Removed in the revised version: C) Control UAS-RasV12, UAS-scrib RNAi, UAS-w RNAi tumors.

Removed in the revised version: D-H) UAS-RasV12, UAS-scrib RNAi tumors co-expressing UAS-RNAi against the PCP genes indicated above each pupa.

Removed in the revised version: B’-H’) Representative third instar larval brain and eye discs for each of the experimental groups from above. Arrowheads in B’ point to the eye-antenna imaginal disc (orange) and the brain lobes (green).

Removed in the revised version: I) Schematic of eye disc RasV12, scrib RNAi, RFP clonal tumors.

Removed in the revised version: J-P) Third instar larval eye discs RFP tumor clones generated via ey-Flp ; FRT42D Gal80 / FRT42D ; act5C>CD2>Gal4, UAS-RFP.

Removed in the revised version: J) Eye discs with non-tumor, control RFP clones.

Removed in the revised version: K) Control RasV12, UAS-scrib RNAi, RFP clonal tumors.

Removed in the revised version: L-P) RasV12, scrib RNAi, RFP clonal tumors carrying the PCP allele indicated above each panel. Scalebars: 100 μm

Earlier version: To study the requirement for PCP signaling in tumors, we expressed RasV12 and scrib RNAi throughout the eye, and simultaneously down-regulated PCP core proteins by also expressing validated RNAi’s against fmi, fz, dsh, vang, and pk. We then imaged tumor size through the pupal cuticle. We observed that tumor growth was unaffected by knockdown of any of the core PCP genes at third instar larval or at early pupal stages (Fig 1A-H). We then asked whether PCP signaling might be required by tumors when they’re confronted by wildtype cells. To test this hypothesis, we created clonal tumors in the eye disc by expressing RasV12 and scrib RNAi under control of ey-Flp. In this model, nearly every cell in the eye disc will perform chromosomal recombination (Fig 1 Suppl 1), producing an eye disc composed of approximately 50% wildtype and 50% tumor cells. Under these conditions, tumor cells expand aggressively, displacing wildtype cells and causing massive overgrowth of the eye disc and metastasis to the larval brain (Fig 1I-K). We created clonal tumors that were also null for either fmi, fz, dgo, vang, or pk and evaluated them for their ability to grow and metastasize when confronted with wildtype cells (Fig 1I-P). Remarkably, only fmi interfered with tumorigenesis, restraining growth of tumors in the eye disc and preventing them from invading neighboring tissues (Fig 1L). To ensure that the growth inhibitory effect of loss of Fmi in clonal tumors but not in whole eye tumors was not a result from comparing a null allele in clonal tumors vs RNAi in whole eye tumors, we generated fmi RNAi clonal tumors and compared them to RNAi control tumors. While the effect was not as dramatic as with null clones, the size of clonal tumors expressing fmi RNAi was dramatically reduced compared to that of tumors expressing w RNAi (Fig 1 Suppl 2A-B). Therefore, the same partial level of fmi knockdown impairs clonal tumor growth but not growth of whole eye tumors. Because this effect was only apparent when tumors were induced in juxtaposition to wildtype cells, we hypothesized that Fmi may only be required in tumors that undergo cell competition.

Revised version: To study the requirement for PCP signaling in tumors, we expressed RasV12 and scrib RNAi throughout the eye, and simultaneously down-regulated PCP core proteins by also expressing validated RNAi’s against fmi, fz, dsh, vang, and pk. We then imaged tumor size through the pupal cuticle. We observed that tumor growth was unaffected by knockdown of any of the core PCP genes at third instar larval or at early pupal stages (Fig 1A-H). We then asked whether PCP signaling might be required by tumors when they’re confronted by wildtype cells. To test this hypothesis, we created clonal tumors in the eye disc by expressing RasV12 and scrib RNAi under control of ey-Flp. In this model, nearly every cell in the eye disc will perform chromosomal recombination (Fig 1 Suppl 1), producing an eye disc composed of approximately 50% wildtype and 50% tumor cells. Under these conditions, tumor cells expand aggressively, displacing wildtype cells and causing massive overgrowth of the eye disc and metastasis to the larval brain (Fig 1I-K). We created clonal tumors that were also null for either fmi, fz, dgo, vang, or pk and evaluated them for their ability to grow and metastasize when confronted with wildtype cells (Fig 1I-P). Remarkably, only fmi interfered with tumorigenesis, restraining growth of tumors in the eye disc and preventing them from invading neighboring tissues (Fig 1L). To ensure that the growth inhibitory effect of loss of Fmi in clonal tumors but not in whole eye tumors was not a result from comparing a null allele in clonal tumors vs RNAi in whole eye tumors, we generated fmi RNAi clonal tumors and compared them to control RNAi clonal tumors. While the effect was not as dramatic as with null clones, the size of fmi RNAi clonal tumors was substantially reduced compared to that of clonal tumors expressing w RNAi (Fig 1 Suppl 2A-B). Therefore, the same partial level of fmi knockdown impairs clonal tumor growth but not growth of whole eye tumors.

Added in the revised version: We also wished to directly compare growth of fmi-null clonal and fmi-null whole eye tumors. Because fmi alleles are embryonic lethal, we generated fmi-null tumors while eliminating WT cells from the eye disc using the GMR-Hid, l(2)CL system (Stowers & Schwarz, 1999). While fmi-null clonal tumors showed restrained growth (Fig 1L), elimination of WT cells in eye discs removed competition, and fmi-null tumors grew to a size comparable to control whole eye tumors (Fig 1 Suppl 2C-D). Because the effect of fmi was only apparent when tumors were induced in juxtaposition to wildtype cells, we hypothesized that Fmi may only be required in tumors that undergo cell competition.

Flamingo is required in winner cells during cell competition

Molecular mechanisms used by tumors in cell competition are sometimes shared by developmental cell competition. Malignant cells have been shown to actively eliminate surrounding cells by mechanical cell competition, apoptosis and engulfment (Levayer et al., 2016; Kohashi et al., 2021), similar to the cell-death mediated elimination of losers during developmental cell competition and EDAC (Maruyama & Fujita, 2017; Parker et al., 2021) Induction of JNK-mediated cell death is also common to developmental competition and tumor cell competition, as shown in intestinal adenomas (Suijkerbuijk et al., 2016). We hypothesized that the requirement of Fmi by competing tumors might also be shared in developmental cell competition. We therefore evaluated the role of Fmi in non-malignant, developmental cell competition models.

Cells expressing high levels of Myc behave as winners when growing among wildtype cells, proliferating faster than wildtype cells and inducing their elimination. Despite Myc being a proto-oncogene, in contrast to RasV12 scrib RNAi tumors, Myc overexpressing clones do not become tumors, but instead expand to occupy larger than normal domains of morphologically normal tissue (Moreno & Basler, 2004; de la Cova et al., 2004; Clavería et al., 2013). As for the RasV12 scrib RNAi tumor model, we used UAS-myc and ey-Flp to create eye discs comprising initially equal populations of Myc overexpressing (>>Myc) and wildtype cells. As expected, >>Myc cells outcompeted wildtype cells, and the majority of cells (∼65%) in late third instar discs were >>Myc cells (Fig 2A). However, when >>Myc cells were depleted of fmi, they failed to outcompete wildtype cells and instead became losers. Wildtype cells outcompeted >>Myc, fmi-/- cells such that late third instar eye discs were composed of only ∼35% >>Myc cells, indicating that wildtype cells behave as winners against >>Myc, fmi-/- cells (Fig 2B,D). We then tested whether Fmi had an effect in the losers during >>Myc competition. Removing Fmi from the WT losers had no effect on this competition, as >>Myc clones did not compete more successfully than when confronting wildtype cells (Fig 3C-D). Importantly, loss of Fmi alone had no effect on competition, as in eye discs comprising initially equal populations of fmi-/- and wildtype cells the fmi mutant cells grew equally well as the wildtype cells (Fig 2 Suppl 1A-C).

Removed in the revised version: Winners require Fmi in developmental supercompetition.

Removed in the revised version: A) Eye imaginal disc clones overexpressing UAS-Myc (>>Myc) and UAS-nGFP under control of the tub>CD2>GAL4 driver. Clones were generated with ey-Flp, so half of the cells are clones and the other half are wildtype twins.

Removed in the revised version: B) Eye imaginal disc >>Myc, UAS-nGFP, fmiE59 clones, competing against wildtype twins.

Removed in the revised version: C) Eye imaginal disc >>Myc, UAS-nGFP clones, competing against fmiE59 twins.

Earlier version: D) Ratio of RFP-labeled clone area versus unlabeled wildtype area. The unlabeled wildtype area was obtained by subtracting the RFP-labeled area from the total eye disc area. A ratio over 1 implies RFP- positive clones are supercompetitors, while a ratio below 1 means the RFP-cells are losers. N = 9 discs (>>Myc), 5 discs (fmiE59, >>Myc), 4 discs (fmiE59 vs >>Myc), groups were analyzed using multiple unpaired, two-tailed t test; p-values: &lt;0.0001 (****), 0.6340 (ns)

Revised version: Winners require Fmi in developmental supercompetition. A) Eye imaginal disc clones overexpressing UAS-Myc (>>Myc) and UAS-nGFP under control of the tub>CD2>GAL4 driver. Clones were generated with ey-Flp, so half of the cells are clones and the other half are wildtype twins. B) Eye imaginal disc >>Myc, UAS-nGFP, fmiE59 clones, competing against wildtype twins. C) Eye imaginal disc >>Myc, UAS-nGFP clones, competing against fmiE59 twins. D) Ratio of RFP-labeled clone area versus unlabeled wildtype area. The unlabeled wildtype area was obtained by subtracting the RFP-labeled area from the total eye disc area. A ratio over 1 implies RFP-positive clones are supercompetitors, while a ratio below 1 means the RFP-cells are losers. N = 9 discs (>>Myc), 5 discs (fmiE59, >>Myc), 4 discs (fmiE59vs >>Myc), groups were analyzed using multiple unpaired, two-tailed t test; p-values: <0.0001 (****), 0.6340 (ns) E) Representative wing imaginal disc overexpressing >>Myc and UAS-nGFP under control of the tub-Gal4 driver. Clones were generated using hsp70-Flp, with a 15 min 37°C heat-shock. wildtype cells are labeled with tub-nRFP. Non-recombinant cells are heterozygous for nRFP, while twin spots are homozygous nRFP. n = 14 discs F) Representative wing imaginal disc with >>Myc, UAS-nGFP, fmiE59 clones, over a wildtype background. Clones were generated in the same fashion as D. n = 19 discs G) Representative wing imaginal disc with >>Myc, UAS-nGFP clones, competing against fmiE59 twin spots. n = 13 discs H) log10 of the >>Myc/Twin spot cell ratio. nGFP total cells were divided by the number of twin spot homozygous nRFP cells, and then log10-transformed. A ratio over 0 indicates supercompetition of nGFP+ clones, and below 0 indicates nGFP+ cells are losers. The difference between groups was analyzed using a one way ANOVA, with Dunnett correction for multiple comparisons; p-value: >>Myc vs fmiE59 >>Myc < 0.0001 (****); >>Myc vs fmiE59 = 0.9705 (ns). Scalebars: 50 μm

Added in the revised version: Winners require fmi in scribble cell competition I-L) scrib RNAi clone analysis in the prospective eye. UAS-scrib RNAi was expressed with the act5C>CD2>Gal4 driver and clones were generated by ey-Flp. Clones are marked both with UAS-nGFP and tub-nRFP. I) Representative disc with control UAS-nGFP clones. J) UAS-scrib RNAi clones versus wildtype twin clones. K) UAS-scrib RNAi, fmiE59 clones versus wildtype twin clones. L) UAS-scrib RNAi clones versus fmiE59 twin clones. This phenotype was lethal. The adult eye shown in L’/L” was from an escaper. Escapers had trouble eclosing and died within hours. I’-L’) Representative adult eye from the genotypes listed above. I”-L”) Fluorescent expression marking the clones in the adult eyes. In the case of scrib RNAi escapers, tub-nRFP fluorescence was barely visible, so the stronger UAS-nGFP is shown. Scalebars: 50 μm

Removed in the revised version: E) Representative wing imaginal disc overexpressing >>Myc and UAS-nGFP under control of the tub-Gal4 driver. Clones were generated using hsp70-Flp, with a 15 min 37°C heat-shock. wildtype cells are labeled with tub-nRFP. Non-recombinant cells are heterozygous for nRFP, while twin spots are homozygous nRFP. n = 14 discs

Removed in the revised version: F) Representative wing imaginal disc with >>Myc, UAS-nGFP, fmiE59 clones, over a wildtype background. Clones were generated in the same fashion as D. n = 19 discs

Removed in the revised version: G) Representative wing imaginal disc with >>Myc, UAS-nGFP clones, competing against fmiE59 twin spots. n = 13 discs

Removed in the revised version: H) log10 of the >>Myc/Twin spot cell ratio. nGFP total cells were divided by the number of twin spot homozygous nRFP cells, and then log10-transformed. A ratio over 0 indicates supercompetition of nGFP+ clones, and below 0 indicates nGFP+ cells are losers. The difference between groups was analyzed using a one way ANOVA, with Dunnett correction for multiple comparisons; p-value: >>Myc vs fmiE59 >>Myc &lt; 0.0001 (****); >>Myc vs fmiE59 = 0.9705 (ns).

Removed in the revised version: Scalebars: 50 μm

Removed in the revised version: Winners require fmi in scribble cell competition

Removed in the revised version: I-L) scrib RNAi clone analysis in the prospective eye. UAS-scrib RNAi was expressed with the act5C>CD2>Gal4 driver and clones were generated by ey-Flp. Clones are marked both with UAS-nGFP and tub-nRFP.

Removed in the revised version: I) Representative disc with control UAS-nGFP clones.

Removed in the revised version: J) UAS-scrib RNAi clones versus wildtype twin clones.

Removed in the revised version: K) UAS-scrib RNAi, fmiE59 clones versus wildtype twin clones.

Removed in the revised version: L) UAS-scrib RNAi clones versus fmiE59 twin clones. This phenotype was lethal. The adult eye shown in L’/L” was from an escaper. Escapers had trouble eclosing and died within hours.

Removed in the revised version: I’-L’) Representative adult eye from the genotypes listed above.

Removed in the revised version: I”-L”) Fluorescent expression marking the clones in the adult eyes. In the case of scrib RNAi escapers, tub-nRFP fluorescence was barely visible, so the stronger UAS-nGFP is shown.

Removed in the revised version: Scalebars: 50 μm

If the requirement of Fmi is a general feature of cell competition, it should be observable in other tissues. We therefore assessed whether depleting Fmi in wing disc >>Myc winner clones would also reverse the competition outcome. We used hsp-Flp to generate twin spot clones expressing UAS-Myc, and as in eye discs, those clones outcompeted neighboring wildtype cells. Myc supercompetitors clones were on average 1.7 times larger than the homozygous RFP+ (hRFP) twin spots (Fig 2E). However, when >>Myc clones lacked Fmi, their ability to compete was severely impaired, producing much smaller clones, being on average half as large as their hRFP twin spots, and thus, like in eye discs, became losers when competing against wildtype cells (Fig. 2F-H). Interestingly, fmiE59Myc clones were also highly fragmented in the wing disc, a phenomenon we did not observe in eye discs. As observed in eye discs, removing Fmi from the wildtype losers had no effect on the ability of >>Myc clones to compete (Fig 3G-H). Similarly, making clones mutant for fmi alone had no effect on their growth in wing discs (Fig 2 Suppl 1D-F).

In the setting of tumorigenesis, either tumor cells or wildtype cells may emerge as winners of cell competition. Tumors outcompete wildtype cells in the processes of invasion and metastasis, whereas transformed pre-malignant cells are often eliminated by wildtype cells through competition-dependent defense mechanisms such as EDAC (Kon et al., 2017; Kon & Fujita, 2021). We therefore further explored the potential requirement for Fmi in a system in which both outcomes occur at different times. In eye discs, cells lacking Scrib display a proto-oncogenic behavior, losing polarity and over-proliferating in the developing eye disc. However, as eye development progresses, they subsequently become losers, and are eliminated via JNK-mediated apoptosis during late third instar and pupal development so that they are virtually absent from the tissue before the fly ecloses (Brumby & Richardson, 2003).

Using the ey-Flp system to activate RNAi against scrib in eye discs, we confirmed that downregulation of scrib generated clones that perdure through larval development (Fig 3D-E,M) but were eliminated from the tissue by the end of pupal development and were not detected in the adult eye (Fig 3F-F’).

Earlier version: Compared to control RFP clones, which represented around 40% of the eye cell population in early third instar larva (Fig 3A-B,M, Suppl 1A) and remained at roughly that proportion in the adult eye (Fig 3C-C’), scrib RNAi clones began to be eliminated at or before the time the morphogenetic furrow progresses (Fig 3E,M, Suppl 1B), and were almost completely absent in adult eyes (Fig 3F-F’). Their elimination left scars in the adult eye, likely because the structure of the eye is established with the passing of the morphogenetic furrow, and the compound eyes were smaller compared to wildtype eyes (Fig 3C,F)).

Revised version: Compared to control RFP clones, which represented around 40% of the eye cell population in early third instar larva (Fig 3A-B, M, Suppl 1A) and remained at roughly that proportion in the adult eye (Fig 3C-C’), scrib RNAi clones began to be eliminated at or before the time the morphogenetic furrow progresses (Fig 3E,M, Suppl 1B), and were almost completely absent in adult eyes (Fig 3F-F’). Their elimination left scars in the adult eye, likely because the structure of the eye is established with the passing of the morphogenetic furrow, and the compound eyes were smaller compared to wildtype eyes (Fig 3C,F)). Wildtype cells therefore behave as winners when confronted with scrib RNAi clones during late larval and pupal development.

Removed in the revised version: Wildtype cells therefore behave as winners when confronted with scrib RNAi clones during late larval and pupal development.

We then evaluated the effect of removing Fmi from the wildtype winner clones in scrib cell competition. Loss of fmi in wildtype cells had little impact on the size of scrib RNAi clones in third instar larval discs (Fig 3G-H,M, Suppl 1C), but as development progressed, wildtype winner cells depleted of fmi lost their ability to compete with and eliminate scrib RNAi clones; the scrib RNAi clones survived during pupal development and indeed constituted the majority of the adult eye (Fig 3I-I’), showing a reversal in the competition outcome much as the outcome of >>Myc competition is reversed when the winners lack Fmi. The surviving scrib RNAi cells failed to differentiate, and instead produced large scars in the adult eye (Fig 3I-I’). We also observed considerable lethality in this population of flies, presumably due to the uncontrolled proliferation of scrib RNAi cells that could not be eliminated.

During >>Myc supercompetition, removing Fmi from the loser wildtype cells showed no effect in cell competition in eye and wing discs (Fig 2C,G). We therefore considered what might happen if scrib RNAi loser clones are made to lack Fmi. However, this situation is more complex than the >>Myc competition, since scrib RNAi clones have been previously documented to initially overproliferate, perhaps behaving as winners, before ultimately becoming losers (Brumby & Richardson, 2003). We were unable to quantify this early proliferation of scrib RNAi eye clones before they are eliminated by their wildtype neighbors, as the discs are too small for us to dissect and count. However, assuming that they are behaving as winners in early larval development, one might expect that removing Fmi from scrib RNAi clones would impair their ability to overproliferate early, such that by third instar, they would be smaller than scrib RNAi clones with intact Fmi. Indeed, this is what we observed (Fig 3J,K,M, Suppl 1D). While some portion of their smaller size relative to control clones is attributed to their elimination by wildtype winners as described above, the remainder may be a result of the loss of their ability to overproliferate earlier in larval development.

The small population of fmiE59 scrib RNAi clones remaining at third instar was almost completely eliminated by the end of larval development, such that the adult eyes were indistinguishable from wildtype eyes (Fig 3L-L’). We hypothesize that the lack of fmi in scrib clones hinders their ability to overproliferate in early larval stages, resulting in smaller clones that are quickly eliminated by their wildtype neighbors, allowing compensatory differentiation to generate a morphologically normal eye.

Loss of fmi triggers cell death and reduces proliferation in “would-be” winner clones

Our observations suggest that lack of fmi renders winner clones and tumors incapable of outcompeting the neighboring tissue in multiple cell competition scenarios. Cell competition relies mainly on two mechanisms to allow winner cells to take over the tissue when confronting less fit cells: faster proliferation than losers and elimination of loser cells by induced apoptosis or extrusion (Amoyel & Bach, 2014; Maruyama & Fujita, 2017; Morata, 2021). We therefore explored how the lack of fmi in tumors and winner >>Myc cells during competition affected both mechanisms.

Earlier version: Drosophila control RasV12, scrib RNAi eye tumors trigger apoptosis in the neighboring cells, as detected by Dcp1 staining (Fig 4A-D), consistent with previous reports (Karim & Rubin, 1998; Pérez et al., 2017) and with their highly invasive behavior when surrounded by wildtype cells. However, when fmi was removed from the clonal tumors, the outcome of this competition was reversed, as described above, and instead, fmi-/-, RasV12, scrib RNAi tumors displayed excess apoptosis compared to the surrounding wildtype tissue (Fig 4E-H). Moreover, we found cell debris from fmi-deficient tumor cells inside lysosomal vesicles in wildtype cells, suggesting that wildtype cells are clearing neighboring loser tumor cells through engulfment (Fig 4 Suppl 1). We then tested whether this reversal of apoptosis burden was specific to the tumor model, or if it is a mechanism shared in other cell competition models. Consistent with the tumor model, eye disc >>Myc cells lacking Fmi displayed widespread apoptosis (Fig 4I-K), likely contributing to the reversal in competition we previously observed (Fig 2A-D).

Revised version: Drosophila control RasV12, scrib RNAi eye tumors trigger apoptosis in the neighboring cells, as detected by Dcp1 staining (Fig 4A-D), consistent with previous reports (Karim & Rubin, 1998; Pérez et al., 2017) and with their highly invasive behavior when surrounded by wildtype cells. However, when fmi was removed from the clonal tumors, the outcome of this competition was reversed, as described above, and instead, fmi-/-, RasV12, scrib RNAi tumors displayed excess apoptosis compared to the surrounding wildtype tissue (Fig 4E-H). Moreover, we found cell debris from fmi-deficient tumor cells inside lysosomal vesicles in wildtype cells, suggesting that wildtype cells are clearing neighboring loser tumor cells through engulfment (Fig 4 Suppl 1). We then tested whether this reversal of apoptosis burden was specific to the tumor model, or if it is a mechanism shared in other cell competition models. Control >>Myc clones showed low levels of apoptosis and similar levels in their neighboring WT cells (Fig. 4I-L; p-value 0.6049), in line with previous results (de la Cova et al., 2004), whereas, consistent with the tumor model, apoptosis was significantly increased in fmiE59, >>Myc clones but not in their WT neighbors (Fig 4M-P, p-value 0.0006), likely contributing to the reversal in competition we previously observed (Fig 2A-D).

Added in the revised version: Lack of Fmi increases cell death and reduces proliferation in would-be winners. A-C) RasV12, scrib RNAi tumors stained for DAPI, Dcp1+ (A), and puc-lacZ (B). C) Merged channels D) Representation of how Dcp1+ staining localizes in the wildtype cells at the boundary with the tumor. E-G) RasV12, scrib RNAi tumors mutant for Fmi, stained with DAPI, Dcp1+ (E), and puc-LacZ (F). G) Merged channels. H) Representation of how Dcp1+ staining localizes in the tumor cells in contact with the surrounding wildtype tissue. Scalebar for A-G: 25 μm I-K) Dcp1+ staining in >>Myc clones in eye discs. >>Myc clones are marked by GFP (I) and were stained against Dcp1+ (J). K) Merged channels, showing apoptotic cells evenly distributed between WT and >>Myc cells. L-N) Dcp1+ staining in >>Myc clones lacking Fmi in the eye disc. Eye disc >>Myc clones are marked by GFP (L) and were stained for Dcp1+ (M). N) Merged channels, showing apoptotic cells localized mainly in the >>Myc, fmiE59 clones. Scalebar: 50 μm O) Quantification of apoptotic cells in WT versus >>Myc clones in eye discs. Apoptosis occurs similarly in WT and >>Myc cells (two-tailed paired t-test; p-value = 0.6049). Points represent the number of apoptotic >>Myc or WT cells per disc (left), or the difference between those values (right; red line indicates mean). P) Quantification of apoptotic cells in WT versus >>Myc, fmiE59 clones in eye discs. Apoptosis is found mainly in >>Myc, fmiE59 cells (two-tailed paired t-test; p-value = 0.0006). Points represent the number of apoptotic >>Myc or WT cells per disc (left), or the difference between those values (right; red line indicates mean). Q-R) Proliferation analysis performed by pHis3 staining in wing discs with either >>Myc clones (Q) or >>Myc, fmiE59 clones (R). Scalebar: 20 μm S) Proliferative ratio of GFP cells in a non-competition Control (n=9 discs), >>Myc (n=9 discs), or >>Myc, fmiE59 (n=13 discs) clones. The proliferative ratio for each group was calculated as the ratio of pHis3 cells within the GFP+ clone versus the non-GFP wildtype tissue and the differences were analyzed as an ordinary ANOVA with a Tukey’s test for multiple comparisons, with all p-values < 0.0001 (****)

Removed in the revised version: Lack of Fmi increases cell death and reduces proliferation in would-be winners.

Removed in the revised version: A-C) RasV12, scrib RNAi tumors stained for DAPI, cleaved Cas3 (A), and puc-lacZ (B). C) Merged channels

Removed in the revised version: D) Representation of how Cas3 staining localizes in the wildtype cells at the boundary with the tumor.

Removed in the revised version: E-G) RasV12, scrib RNAi tumors mutant for Fmi, stained with DAPI, cleaved Cas3 (E), and puc-LacZ (F). G) Merged channels.

Removed in the revised version: H) Representation of how Cas3 staining localizes in the tumor cells in contact with the surrounding wildtype tissue.

Removed in the revised version: Scalebar for A-G: 25 μm

Removed in the revised version: I-K) Cleaved Cas3 staining in >>Myc clones lacking Fmi in the eye disc. Eye disc Myc clones are marked by GFP (I) and were stained for cleaved Cas3 (J). K) Merged channels, showing apoptotic cells localized only in the Myc, fmiE59 clones.

Removed in the revised version: Scalebar: 50 μm

Removed in the revised version: L-M) Proliferation analysis performed by pHis3 staining in wing discs with either >>Myc clones (L) or >>Myc, fmiE59 clones (M). Scalebar: 20 μm

Removed in the revised version: N) Proliferative ratio of GFP cells in a non-competition Control (n=9 discs), >>Myc (n=9 discs), or >>Myc, fmiE59 (n=13 discs) clones. The proliferative ratio for each group was calculated as the ratio of pHis3 cells within the GFP+ clone versus the non-GFP wildtype tissue and the differences were analyzed as an ordinary ANOVA with a Tukey’s test for multiple comparisons, with all p-values &lt; 0.0001 (****)

Earlier version: JNK signaling plays a prominent role during cell competition. Previous reports have shown that activation of JNK mediates engulfment and elimination of scribble clones during cell competition (Ohsawa et al., 2011; Yamamoto et al., 2017). However, JNK signaling is also responsible for increased tumor growth and invasion in Ras-activated cells or scrib mutants (Igaki et al., 2006; Uhlirova & Bohmann, 2006; Leong et al., 2009). We therefore asked whether the role of Fmi in tumorigenesis and cell competition could be related to the regulation of JNK signaling. Confirming previous observations, we detect activation of JNK signaling via the puckered lacZ reporter pucE69 both in RasV12, scrib RNAi tumors (Fig 4A-C), and scrib RNAi clones in eye discs (Figure 4 Suppl 2A-C). However, RasV12 tumors lacking fmi, despite displaying cell death in tumor cells competing with wildtype neighbors, show no change in Puc activation (Fig 4E-G). Similarly, scrib RNAi clones activate JNK signaling independently of fmi (Fig 4 Suppl 2D-F). We therefore conclude that Fmi acts independently of JNK signaling.

Revised version: JNK signaling plays a prominent role during cell competition. Previous reports have shown that activation of JNK mediates engulfment and elimination of scribble clones during cell competition (Ohsawa et al., 2011; Yamamoto et al., 2017). However, JNK signaling is also responsible for increased tumor growth and invasion in Ras-activated cells or scrib mutants (Igaki et al., 2006; Uhlirova & Bohmann, 2006; Leong et al., 2009). We therefore asked whether the role of Fmi in tumorigenesis and cell competition could be related to the regulation of JNK signaling. Confirming previous observations, we detect activation of JNK signaling via the puckered lacZ reporter pucE69 both in RasV12, scrib RNAi tumors (Fig 4A-C), and scrib RNAi clones in eye discs (Fig 4 Suppl 2A-C). However, RasV12 tumors lacking fmi, despite displaying cell death in tumor cells competing with wildtype neighbors, show no change in Puc activation (Fig 4E-G). Similarly, scrib RNAi clones activate JNK signaling independently of fmi (Fig 4 Suppl 2D-F). We therefore conclude that Fmi acts independently of JNK signaling.

Earlier version: Cell competition does not rely solely on apoptosis to eliminate loser cells. An increased proliferation rate of winners is also a hallmark of cell competition (Morata, 2021; Madan et al., 2022). To evaluate whether Fmi is also required to maintain a higher proliferative ratio in winners, we quantified mitotic cells in winner clones with and without Fmi. The wing disc displays distributed cell divisions throughout larval development, in contrast to eye discs, where passing of the morphogenetic furrow limits proliferation as cells start differentiating into compound eye cells (Wolff & Ready, 1991; Tsachaki & Sprecher, 2012). Therefore, for the quantification of proliferation, the wing is a better model than the eye disc. When wing disc >>Myc clones were depleted of Fmi, cell proliferation, as measured by pHis3 staining, was significantly reduced (Fig 4L-N). Taken together, these observations directly link the need for Fmi to proliferation and induced apoptosis, the two key events of cell competition, in several models of cell competition.

Revised version: Cell competition does not rely solely on apoptosis to eliminate loser cells. An increased proliferation rate of winners is also a hallmark of cell competition (Morata, 2021; Madan et al., 2022). To evaluate whether Fmi is also required to maintain a higher proliferative ratio in winners, we quantified mitotic cells in winner clones with and without Fmi. The wing disc displays distributed cell divisions throughout larval development, in contrast to eye discs, where passing of the morphogenetic furrow limits proliferation as cells start differentiating into compound eye cells (Wolff & Ready, 1991; Tsachaki & Sprecher, 2012). Therefore, for the quantification of proliferation, the wing is a better model than the eye disc. When wing disc >>Myc clones were depleted of Fmi, cell proliferation, as measured by pHis3 staining, was significantly reduced (Fig 4Q-S). Taken together, these observations directly link the need for Fmi to proliferation and induced apoptosis, the two key events of cell competition, in several models of cell competition.

Fmi mediated cell-cell communication is not required for competition between winners and losers

Added in the revised version: Previously, some key players involved in cell competition, including Flower (Rhiner et al., 2010) and Xrp-1 (Baillon et al., 2018), were shown to be either up- or down-regulated during competition. We evaluated whether Fmi protein levels might also be affected during competition. Fmi is ubiquitously expressed at low levels in Drosophila larva, pupa and adult (Brown et al., 2014). We stained third instar larval wing discs with >>Myc clones to determine whether >>Myc supercompetition affects the levels of Fmi protein at the membrane, either inside the clone or near the boundary where cell competition occurs. We observed no change in Fmi protein levels (Fig 5 Suppl 1), suggesting that Fmi is involved in competition through a mechanism that does not rely on altering protein levels.

Added in the revised version: Vang is not required for >>Myc supercompetition A) Representative wing imaginal disc showing the GFP-tagged, vangA3 >>Myc clones. Clones were generated using hsp70-Flp B) RFP-labeled twin spots for the clones shown in A. Twin spot homozygous for RFP can be observed adjacent to the supercompetitor clones. C) Graph showing the total GFP and homozygous RFP counts per disc, with each disc counts linked by a line and the log10 ratio of GFP/hRFP cells per disc (n=12 discs). Differences between GFP/hRFP clone size were analyzed using a two-tailed, ratio-paired t test. p-value<0.0001. When the ratio was compared with a two-tailed t test to >>Myc, the difference was not significant, with a p-value of 0.078. Scalebar: 50 μm

Cell competition is thought to rely on intercellular communication to compare fitness and determine the outcome between prospective winner and loser cells (Kon et al., 2017; Baker, 2020; Ogawa et al., 2021). If Fmi is involved in these determinative intercellular communication events by signaling as a trans-homodimer, it should be required in both prospective winner and loser cells. Our previous eye and wing disc >>Myc supercompetition results already suggested this is not the case. While the removal of fmi caused winner clones to effectively become losers during >>Myc competition (Fig 2), removing fmi from the losers had no effect on the losers’ outcome, either in eye discs (Fig 2C-D) or wing discs (Fig 2G-H). In both tissues, winner >>Myc clones showed no change in relative size to their twin spot counterparts, nor were loser clones eliminated more effectively (Fig 2D and 2H).

Our results so far suggested that the effect Fmi exerts on cell competition does not operate through bidirectional intercellular PCP communication. To further consolidate this hypothesis, we examined the effect of removing a dedicated core PCP protein other than Fmi from >>Myc supercompetitors.

Generating vangA3 >>Myc clones in the wing disc, we observed that the >>Myc clones were unaffected and remained winners (Fig 5). vangA3, >>Myc clones were on average 2.2±0.82 times larger than their hRFP twins, not significantly different from the 1.7±0.45 fold value for >>Myc clones (Fig 2H; Fig 5C).

Removed in the revised version: Vang is not required for >>Myc supercompetition

Removed in the revised version: A) Representative wing imaginal disc showing the GFP-tagged, vangA3 >>Myc clones. Clones were generated using hsp70-Flp

Removed in the revised version: B) RFP-labeled twin spots for the clones shown in A. Twin spot homozygous for RFP can be observed adjacent to the supercompetitor clones.

Removed in the revised version: C) Graph showing the total GFP and homozygous RFP counts per disc, with each disc counts linked by a line and the log10 ratio of GFP/hRFP cells per disc (n=12 discs). Differences between GFP/hRFP clone size were analyzed using a two-tailed, ratio-paired t test. p-value&lt;0.0001. When the ratio was compared with a two-tailed t test to >>Myc, the difference was not significant, with a p-value of 0.078.

Removed in the revised version: Scalebar: 50 μm

The activity of Fmi in cell competition does not require its cadherin domains

Earlier version: Our results have so far demonstrated that Fmi is required only in winner cells and that its function in cell competition is independent of PCP. Furthermore, our results rule out a role for Fmi homodimers in directly communicating fitness information between prospective winner and loser cells. However, the possibility remained that Fmi might function through trans-homodimerization between prospective winner cells to sustain winner status via a signal, via adhesion, or both. To ask if Fmi fulfills its role by mediating adhesion, we tested whether we could restore winner status to >>Myc clones that lack Fmi by providing a transgenic fmi that lacks the 9 cadherin domains of Fmi (arm-fmiCad).

Revised version: Our results have so far demonstrated that Fmi is required only in winner cells and that its function in cell competition is independent of PCP. Furthermore, our results rule out a role for Fmi homodimers in directly communicating fitness information between prospective winner and loser cells. However, the possibility remained that Fmi might function through trans-homodimerization between prospective winner cells to sustain winner status via a signal, via adhesion, or both. To ask if Fmi fulfills its role by mediating adhesion, we tested whether we could restore winner status to >>Myc clones that lack Fmi by providing a transgenic fmi that lacks the 9 cadherin domains of Fmi (arm-fmi△iCad).

Coexpression of arm-fmiΔCad in fmiE59, >>Myc clones re-established the ability of these clones to outcompete their neighbors, and they again became supercompetitors (Fig. 6). When comparing the ratio of >>Myc clone cell number to twin spot cell number, the presence of arm-fmiΔcad restores competition to a level comparable to >>Myc supercompetitors (1.43±0.48 GFP/hRFP cells vs 1.71±0.45 GFP/hRFP cells respectively), far above the 0.49±0.29 value for fmiE59 >>Myc clones compared to twin spots.

Removed in the revised version: The Fmi cadherin repeats are not required for cell competition.

Removed in the revised version: A) Representative disc with nGFP-labeled, >>Myc (A), >>Myc, fmiE59 (B) or >>Myc, fmiE59, arm-fmiΔcad

Removed in the revised version: (C) clones competing against wildtype twin spots in the wing disc. Twin spot clones are labeled with homozygous nRFP and clones were generated with hsp70-Flp.

Earlier version: D) Ratio of GFP versus RFP cells in the three groups, represented as the log10(GFP/hRFP) cell ratios. To evaluate the effect of the arm-fmiΔcad rescue (n=14 discs), the GFP/hRFP cell ratio was directly compared against the two other groups, already quantified and shown in Fig. 2F. Differences between the groups were analyzed using an unpaired, ordinary one-way ANOVA, which found a p-value &lt;0.0001. Inter-group differences were analyzed with a Tukey multiple comparisons test, which found no differences between >>Myc and >>Myc, fmiE59, arm-fmiΔcad clones, whereas both groups strongly differed from >>Myc, fmiE59 clones, both returning a p-value &lt;0.0001 (****) when compared directly against >>Myc, fmiE59.

Revised version: The Fmi cadherin repeats are not required for cell competition. A) Representative disc with nGFP-labeled, >>Myc (A), >>Myc, fmiE59 (B) or >>Myc, fmiE59, arm-fmiΔcad (C) clones competing against wildtype twin spots in the wing disc. Twin spot clones are labeled with homozygous nRFP and clones were generated with hsp70-Flp. D) Ratio of GFP versus RFP cells in the three groups, represented as the log10(GFP/hRFP) cell ratios. To evaluate the effect of the arm-fmiΔcad rescue (n=14 discs), the GFP/hRFP cell ratio was directly compared against the two other groups, already quantified and shown in Fig. 2F. Differences between the groups were analyzed using an unpaired, ordinary one-way ANOVA, which found a p-value <0.0001. Inter-group differences were analyzed with a Tukey multiple comparisons test, which found no differences between >>Myc and >>Myc, fmiE59, arm-fmiΔcad clones, whereas both groups strongly differed from >>Myc, fmiE59 clones, both returning a p-value <0.0001 (****) when compared directly against >>Myc, fmiE59. Scalebar: 50 μm

Removed in the revised version: Scalebar: 50 μm

Earlier version: Supercompetitor clones mutant for fmi consistently show clone fragmentation (Figs. 2F, 4M, 6B). Despite lacking the cadherin repeats, rescued clones not only restored their supercompetitor status, but fully recovered the ability of >>Myc clones to remain cohesive (Fig. 6C). This suggests that the clone fragmentation we observed is unlikely due to the ability of Fmi to contribute to adhesion and more likely caused by a feature of the competition between wildtype cells and fmi-/- >>Myc loser clones that causes their elimination.

Revised version: Supercompetitor clones mutant for fmi consistently show clone fragmentation (Figs. 2F, 4R, 6B). Despite lacking the cadherin repeats, rescued clones not only restored their supercompetitor status, but fully recovered the ability of >>Myc clones to remain cohesive (Fig. 6C). This suggests that the clone fragmentation we observed is unlikely due to the ability of Fmi to contribute to adhesion and more likely caused by a feature of the competition between wildtype cells and fmi-/- >>Myc loser clones that causes their elimination.

Taken together, these results demonstrate that, while Fmi is essential in winner cells to eliminate less fit neighbors, this effect is independent of PCP or other homodimer-mediated signaling, and independent of Fmi-mediated cell adhesion, suggesting instead an as yet uncharacterized function for Fmi.

Discussion

We have identified a requirement for Fmi in winner cells in both tumorigenic and developmental cell competition models. Cells that would otherwise behave as winners instead behave as losers when they lack Fmi. Fmi is notable in that it is required in cell competition in each of four distinct competition scenarios examined: RasV12 scrib RNAi tumors, Myc supercompetitor clones in eye and wing discs, wildtype cells vs scrib RNAi loser clones in pupal eyes, and likely scrib RNAi clones in larval eye discs. Just how universal this requirement is in other cell competition scenarios in Drosophila and perhaps in competition in other organisms remains to be determined.

Fmi acts in cell competition independently of PCP

Several arguments support the conclusion that the role for Fmi in cell competition is distinct from its role in PCP signaling. First, Fmi is the only core PCP component among the six that were surveyed to inhibit the ability of RasV12 scrib RNAi tumor clones to compete in the eye. If Fmi’s role in cell competition were to signal winner fate to losers and vice versa by mirroring its role in PCP signaling, where it signals the presence of proximal (Vang, Pk) components in one direction and the presence of distal components (Fz, Dsh, Dgo) in the other direction between adjacent cells (Lawrence et al., 2004; Strutt & Strutt, 2007; Strutt & Strutt, 2008; Chen et al., 2008; Struhl et al., 2012), then one might expect either the proximal or distal components to also be required in winner clones. No such requirement was observed. Second, in PCP signaling, Fmi functions as a trans-homodimer to transmit those signals and requires the cadherin repeats and other extracellular domains, implying its function as a trans-homodimer (Kimura et al., 2006). In PCP signaling, removing Fmi from either of two adjacent cells completely blocks PCP signaling. In contrast, in cell competition, while Fmi is required in winners, removing Fmi from losers has no effect on competition. Thus, the model of bidirectional signaling via Fmi is not supported by our results.

These observations, however, do not rule out the possibility that Fmi trans-homodimers contribute to intercellular signaling among winner cells. Nonetheless, as discussed above, adhesion seems not to be a meaningful part of the function for Fmi in winners, as an adhesion deficient Fmi construct (fmiΔCad) fully rescues both competition and clone adhesion. The potential contributions of adhesion vs other possible signaling mechanisms are discussed at more length below.

Unlike the other core PCP genes, fmi-/- mutations are lethal due to requirements in the nervous system (Usui et al., 1999). Though not fully characterized, it’s roles in the nervous system appear to be distinct from PCP signaling. Fmi is required for outgrowth and guidance of the R8 axon of the eye to the M3 layer of the medulla via a mechanism that appears independent of other components of the core PCP signaling pathway (Gao et al., 2000; Lee et al., 2003; Senti et al., 2003), but does interact with Golden goal, a transmembrane phosphoprotein that is not associated with PCP signaling (Takechi et al., 2021). Growth of the dendrites of dorsal da neurons is also regulated by Fmi. During embryogenesis, da dendrites in fmi mutant embryos emerge precociously and overgrow as they approach the dorsal midline, and later, during larval growth, dendrites from opposite sides fail to avoid each other (tile) and instead overlap (Gao et al., 2000; Sweeney et al., 2002; Kimura et al., 2006).

Fmi is classified as an atypical cadherin and a Class-B adhesion G Protein-Coupled Receptor (AGPCR), as it contains in its extracellular domain several conserved functional domains including cadherin repeats, epidermal growth factor-like repeats, laminin A G-type repeats, and a GPCR autoproteolytic inducing (GAIN) domain that contains within it a GPCR proteolytic site (GPS) (Rosa et al., 2021; Einspahr & Tilley, 2022; Sreepada et al., 2022). These extracellular domains are followed by seven transmembrane domains and an intracellular C-terminal domain. Although much remains to be learned about this large subfamily of GPCRs, a general model has emerged in which activation by membrane bound or extracellular protein, peptide, proteoglycan or small molecule ligand, or mechanical force exposes a tethered ligand at the N-terminus of the C-terminal fragment in the GAIN domain that, upon exposure, interacts with the transmembrane portion to activate a G-protein signaling cascade. In many but not all AGPCRs, the tethered ligand is exposed by GAIN domain-mediated autoprotolysis of its GPS. Non-cleaved AGPCRs are hypothesized to expose the tethered ligand by an allosteric conformational change. Some de-orphanized AGPCRs interact with multiple ligands, and ligand binding can result in partial or full activation. In some cases, it appears that engineered truncation of portions of the extracellular domain can produce some level of ligand independent activation (Rosa et al., 2021).

Earlier version: When expressed in da neurons, a Fmi construct lacking the cadherin, laminin G and EGF-like repeats partially rescued the embryonic dendritic overgrowth phenotype (Kimura et al., 2006), suggesting a function independent of homodimerization. The G protein Gαq (Gq) has been proposed to function downstream of Fmi to mediate this repressive function (Wang et al., 2016). A recent preprint reports the resolved structure of the CELSR1 extracellular domain, showing that the protein has two distinct domains: an adhesion domain comprised of the first 8 Cadherin repeats, and a compact domain that extends from the ninth cadherin repeat to the transmembrane domains that is involved in GPCR signaling. Indeed, they demonstrated that a CELSR1 construct lacking only the Cadherin repeats 1-8 retains the ability to activate Gαs, which has been predicted to interact with CELSR1 (Bandekar et al., 2024). Notably, our results showed that a similar Fmi construct lacking the cadherin domains substantially rescues the requirement for Fmi in winner cells during cell competition (Fig. 6). These observations suggest that supplying adhesion is not the principal function of Fmi in these events, but are consistent with the possibility that homodimeric adhesion, or interaction with a different ligand, normally activates the receptor, and that the truncated FmiΔCad behaves as a constitutively activated receptor capable of binding to either Gα proteins. While no biochemical characterization of Fmi has been reported, the human orthologs CELSR1-3 have been studied in detail. CELSR2 is autoproteolytically cleaved while CELSRs 1 and 3 are not, yet all three couple to GαS (Huong Bui et al., 2023). Additional efforts will be required to determine the functional ligand(s) for Fmi and whether it signals similarly. Furthermore, AGPCRs participate in a wide variety of developmental and physiologic events through diverse effectors (Einspahr & Tilley, 2022; Sreepada et al., 2022). The pathway by which Fmi participates in cell competition remains to be explored.

Revised version: When expressed in da neurons, a Fmi construct lacking the cadherin, laminin G and EGF-like repeats partially rescued the embryonic dendritic overgrowth phenotype (Kimura et al., 2006), suggesting a function independent of homodimerization. The G protein Gαq (Gq) has been proposed to function downstream of Fmi to mediate this repressive function (Wang et al., 2016). A recent preprint reports the resolved structure of the CELSR1 extracellular domain, showing that the protein has two distinct domains: an adhesion domain comprised of the first 8 Cadherin repeats, and a compact domain that extends from the ninth cadherin repeat to the transmembrane domains that is involved in GPCR signaling. Indeed, they demonstrated that a CELSR1 construct lacking only the Cadherin repeats 1-8 retains the ability to activate Gαs, which has been predicted to interact with CELSR1 (Bandekar et al., 2024). Notably, our results showed that a similar Fmi construct lacking the cadherin domains substantially rescues the requirement for Fmi in winner cells during cell competition (Fig. 6). These observations suggest that supplying adhesion is not the principal function of Fmi in these events, but are consistent with the possibility that homodimeric adhesion, or interaction with a different ligand, normally activates the receptor, and that the truncated FmiΔCad behaves as a constitutively activated receptor capable of binding to either Gα proteins. While no biochemical characterization of Fmi has been reported, the human orthologs CELSR1-3 have been studied in detail. CELSR2 is autoproteolytically cleaved while CELSRs 1 and 3 are not, yet all three couple to GαS (Huong Bui et al., 2023). Additional efforts will be required to determine the functional ligand(s) for Fmi and whether it signals similarly.

Added in the revised version: Furthermore, AGPCRs participate in a wide variety of developmental and physiologic events through diverse effectors (Einspahr & Tilley, 2022; Sreepada et al., 2022). The pathway by which Fmi participates in cell competition remains to be explored.

Added in the revised version: Fmi, cell competition and cancer

Earlier version: Fmi, cell competition and cancer When the first examples of super-competition were observed in Myc clones and the Hippo pathway (Moreno & Basler, 2004; de la Cova et al., 2004; Harvey & Tapon, 2007), they hinted at the possibility that tumors, which behave like super-competitors, could use similar mechanisms to outcompete wildtype cells. Understanding tumor competition may open new avenues for early detection and therapy (Baker & Li, 2008; Moreno, 2008).

Revised version: When the first examples of super-competition were observed in Myc clones and the Hippo pathway (Moreno & Basler, 2004; de la Cova et al., 2004; Harvey & Tapon, 2007), they hinted at the possibility that tumors, which behave like super-competitors, could use similar mechanisms to outcompete wildtype cells. Understanding tumor competition may open new avenues for early detection and therapy (Baker & Li, 2008; Moreno, 2008).

Research in Drosophila and mammals has shown that cell competition plays a dual role during tumorigenesis. Cells harboring mutations in proto-oncogenes or tumor-suppressor genes often behave as losers (Maruyama & Fujita, 2017; Morata & Calleja, 2020; Kanda & Igaki, 2020). Through the process of EDAC, epithelial tissues use cell competition to eliminate transformed pre-neoplastic cells by removing them from the tissue via directed cell death or extrusion (Kon et al., 2017; Watanabe et al., 2018). Pre-neoplastic cells that escape EDAC may accumulate additional mutations to become malignant tumors (Watanabe et al., 2018). Malignant tumors not only escape EDAC, but acquire properties that allow them to outcompete wildtype cells, facilitating invasion and metastasis (Suijkerbuijk et al., 2016; Kohashi et al., 2021). Furthermore, competition between clones within tumors further selects for more aggressive tumor behavior (Parker et al., 2021).

Another commonality between developmental and oncogenic cell competition is the involvement of the transmembrane protein Flower (Fwe). In both Drosophila and mammals, multiple isoforms of Fwe signal fitness; expression of FweLose isoforms mark losers for elimination (Rhiner et al., 2010; Merino et al., 2013; Levayer et al., 2015; Madan et al., 2019). Forced expression of FweLose induces cell competition and elimination of the loser, suggesting that Fwe comparison is involved in the sensing and/or initiation of differential fitness. This contrasts with Fmi, whose differential expression does not act as a trigger for competition, but which is needed in winners to allow them to win, suggesting that Fmi is involved after sensing in the execution of functions necessary to manifest winner behavior.

Evidence is accumulating that a human ortholog of Fmi, CELSR3, is expressed at high levels in a range of solid tumors, including lung, prostate, pancreatic, hepatic, ovarian and colorectal cancers, and in some cases has been shown to be associated with poor prognosis (Katoh & Katoh, 2007; Erkan et al., 2010; Asad et al., 2014; Goryca et al., 2018; Li et al., 2021; Chen et al., 2021a). Recently, CELSR1 upregulation has also been linked to poor ovarian cancer prognosis, likely by promoting proliferation, migration, and invasion (Zuo et al., 2023). If CELSR1/3 are promoting winner cell behavior in these tumors, as might be predicted from its function in Drosophila, this could provide the rationale for future efforts to understand the mechanism by which Fmi/CELSR3 facilitates cell competition, with the goal of identifying an intervention that could blunt or perhaps even eliminate the aggressiveness of an array of highly morbid cancers.